Imidacloprid disrupts larval molting regulation and nutrient energy metabolism, causing developmental delay in honey bee Apis mellifera

Author Response

The following is the authors’ response to the original reviews.

We sincerely appreciate the opportunity to revise the manuscript and the reviewers' critical comments and valuable suggestions. After carefully revising the manuscript, we strongly believe that the reviewers' comments are invaluable and will significantly enhance the quality of the manuscript and contribute to our future research. Following the reviewers' comments, we conducted a comprehensive and meticulous review, addressing each point individually and making extensive modifications and corrections. The responses to each question are provided in a point-by-point manner as follows:

Reviewer 1:

This study delves into the impact of imidacloprid, an insecticide documented for its toxicity towards honeybees, on the development of bee larvae. The investigation involved exposing bee larvae to various concentrations of imidacloprid, and observing the resultant effects.

The findings of this study revealed that imidacloprid exerted a dose-dependent delay in the development of bee larvae, marked by reductions in body mass, width, and an overall decline in the growth index. Moreover, at elevated concentrations, imidacloprid was observed to impair neural transmission, induce oxidative stress, inflict damage to the gut, and inhibit hormones and genes essential for development. The larvae were found to engage antioxidant defense systems and deploy detoxification mechanisms to mitigate these effects.

However, the manuscript could be significantly enhanced through several improvements. Firstly, the structure of the manuscript warrants refinement to foster coherence and clarity. Additionally, there is a need for careful reevaluation of the concentrations of imidacloprid employed in the study, to ensure their relevance and applicability. In terms of references, greater attention to accuracy in citation is imperative.

Furthermore, while the authors have provided an overview of the general effects of imidacloprid on both vertebrates and invertebrates, the inclusion of a more exhaustive literature review with a specific focus on honey bees and other insects would bolster the context and significance of this research. This would be particularly beneficial in the introduction section, which should be subjected to a major revision.

In summary, this study offers preliminary evidence of the detrimental effects of imidacloprid on the development of bee larvae by interfering with molting and metabolism. This research holds potential as a valuable resource for assessing the risks posed by pesticides to juvenile stages of various animal species.

On behalf of all the authors, I express our most sincere gratitude for your critical comments and suggestions. Following your suggestions, we have thoroughly reviewed and revised the entire manuscript, including the issues of imidacloprid concentration and citation accuracy you raised. More importantly, we have significantly revised the structure and content of the introductory section of the manuscript to include many more detailed reviews of critical literature, with particular attention to the overview of relevant research on honey bees, Drosophila, and other insects, to promote coherence and clarity of the introduction and to enhance the context and importance of this research. We hope that these changes meet with your approval. Overall, your valuable comments have greatly improved the quality of the manuscript and will facilitate our future research.

Q1: Line 48, "Adults exposed to high doses of imidacloprid experience", please provide a more precise value for the high doses.

Thank you very much for your comments. Following your suggestion, we have provided precise values for high doses of imidacloprid for adult exposure based on the study by Dr. Wu et al. 2001.

Q2: Line 82, There are several larvae effect reports using next generation sequencing approach. The authors should include those related references in this section.

Thank you for your comments. We have included relevant references in our revised manuscript.

Q3: Line 394, for the concentration design, the maximum concentration of imidacloprid used in this study is 377 ppb, which is from the imidacloprid residue level in beeswax. Bees don't consume beeswax, and the reference is wrong.

Thank you for raising this critical issue. As you point out, bees do not eat beeswax, but it is important to stress that this may well mean that the bee larvae themselves are exposed to higher doses. Therefore, in this study, we ultimately designed for the worst-case scenario of 377 ppb of imidacloprid residues in beeswax. We would like your agreement on this point. In addition, we have corrected the citation errors in the references here and included them in the revised manuscript.

Reviewer 2:

This study provides evidence on the ability of sublethal imidacloprid doses to affect growth and development of honeybee larva. While checking the effect of doses that do not impact survival or food intake, the authors found changes in the expression of genes related to energy metabolism, antioxidant response, and P450 metabolism. The authors also identified cell death in the alimentary canal, and disturbances in levels of ROS markers, molting hormones, weight and growth ratio. The study strengths come from applying these different approaches to investigate the impacts of imidacloprid exposure. The study weaknesses are not providing an in-depth investigation of the mechanisms behind the impacts observed and not bringing the results in light of the current literature. For instance, the authors' hypothesis is based on two main points, the generation of ROS that leads to gut cell death and energy dysfunction, and the increased P450 expression. They propose this increases P450 expression which in turn increases energy consumption and could contribute to developmental retardation. There is however no investigation on the mechanisms of ROS generation (it could be through mitochondrial damage, Nox/ Duox activity, NOS activity, P450s activity, etc). A link between higher P450 expression and increased energy consumption leading to energy deprivation is also missing. It would also be important for the authors to provide a more complete literature review as previous works have investigated imidacloprid sublethal dose impacts in larval stages for bees and other insect models.

I greatly appreciate your insightful comments and valuable suggestions on behalf of all the authors. Thank you for identifying the limitations of this study and providing valuable comments and suggestions. These comments and suggestions have significantly improved the quality of the paper and will facilitate our future research. Following your comments, we have revised and corrected the manuscript point by point. We hope that these corrections meet with your approval.

Q1: Abstract: It would be important to rephrase the abstract to make it clear when authors are talking about gene expression results or functional assays.

Thank you for your comment. Following your suggestion, we have revised the abstract to make it clearer, especially the description of the gene expression results. Please see lines 15-34 in our revised manuscript.

“Abstract Imidacloprid is a global health threat that severely poisons the economically and ecologically important honeybee pollinator, Apis mellifera. However, its effects on developing bee larvae remain largely unexplored. Our pilot study showed that imidacloprid causes developmental delay in bee larvae, but the underlying toxicological mechanisms remain incompletely understood. In this study, we exposed bee larvae to imidacloprid at environmentally relevant concentrations of 0.7, 1.2, 3.1, and 377 ppb. There was a marked dose-dependent delay in larval development, characterized by reductions in body mass, width, and growth index. However, imidacloprid did not affect larval survival and food consumption. The primary toxicological effects induced by elevated concentrations of imidacloprid (377 ppb) included inhibition of neural transmission gene expression, induction of oxidative stress, gut structural damage, and apoptosis, inhibition of developmental regulatory hormones and genes, suppression of gene expression levels involved in proteolysis, amino acid transport, protein synthesis, carbohydrate catabolism, oxidative phosphorylation, and glycolysis energy production. In addition, we found that the larvae may use antioxidant defenses and P450 detoxification mechanisms to mitigate the effects of imidacloprid. Ultimately, this study provides the first evidence that environmentally exposed imidacloprid can affect the growth and development of bee larvae by disrupting molting regulation and limiting the metabolism and utilization of dietary nutrients and energy. These findings have broader implications for studies assessing pesticide hazards in other juvenile animals”

Q2: Line 55-58: rephrase the sentences to make it clear that imidacloprid was not created in 1925, but only in the 90's.

Thank you for pointing out this error. We have corrected the citation. Please see the line 58 in our revised version.

Q3: Line 88: typo: " remain to be systematically investigated"

Thank you for pointing out this error. We have rewritten the sentence. Please see lines 121-122 in our revised manuscript.

Q4: Introduction is lacking important citations, a few of the important ones are: Farooqui 2013 (doi: 10.1016/j.neuint.2012.09.020.) - hypothesis linking neonic exposure, nAChRs receptors, and ROS in honeybees; Ihara et al 2020 (https://doi.org/10.1073/pnas.2003667117) - the targets of imidacloprid in honeybees; Martelli et al 2020 (https://doi.org/10.1073/pnas.2011828117) - mechanistic investigation of imidacloprid sublethal damage in Drosophila; Whitehorn et al 2018 (doi: 10.7717/peerj.4772) - investigation of imidacloprid sublethal dose impact on growth and development of butterflies; Chen et al 2021 (doi: 10.3390/ijms222111835) - sublethal effects of imidacloprid exposure on gene expression in honeybees at different life stages. It is important that the authors perform a more complete literature search to compare their work to previous ones, drawing conclusions and highlighting their novelties.

We greatly appreciate your insightful comments and valuable suggestions. Following your suggestions, we have made significant revisions to the structure and content of the Introduction section. We have incorporated the critical literature you provided and other relevant literature reviews, with a particular emphasis on studies of bees, fruit flies, and other insects. These revisions aim to improve the coherence, clarity, background, and significance of the Introduction. We hope that these modifications meet with your approval. Please see the red text in the Introduction section in our revised version.

Q5: Line 104: Explanation on the doses used should be included here, not later in the methods. Also, important to highlight that whereas the doses tested were found in bee products, they likely mean that the bees themselves were being exposed to even higher doses.

Thank you for your comment. Following your suggestions, we have moved the explanation of the imidacloprid doses used in this study to the Results section, as you mentioned. Please see lines 138-142 in our revised manuscript.

Q6: Line 112: It is important to identify the neuronal targets of imidacloprid in honeybees. Many are known. Some of the nAChRs targets were not investigated in this study (such as subunit alpha8 and beta1). Plus, is alpha2 an imidacloprid target? How does the expression of other nAChRs subunits compares? Importantly, these genes are expressed mostly in the nervous system, so a more correct approach would be a tissue specific analysis. The lack of tissue specific analysis is a consistent flaw throughout the methodological design.

Thank you very much for your important comment. Bees have more than ten nAChR subunit members. Imidacloprid inhibits acetylcholinesterase activity by competitively binding to acetylcholinesterase receptors. As you noted, this study did not investigate the expression of all nAChR subunits, including the alpha8 and beta1 subunits, in different tissues, which is a shortcoming of our study. We have always failed to make a technological breakthrough and cannot dissect to obtain important tissues from developing larvae alone. We have therefore had to abandon this design and use the whole larva as a sample for measurement. We are aware that this is a shortcoming of this research. In the future, we will make a breakthrough in technology and conduct a comparative analysis of all nAChR subunit genes in different organizations and developmental stages to obtain more comprehensive and accurate data. Thank you again for raising this important issue and for your valuable suggestions.

Q7 ~ Q9: Line 125: P450s expression may have opposite behavior when exposed to insecticides depending on tissue (such as brain and fat body). When checking whole larva gene expression, the tissue specific profiles become diluted and thus less reliable (for reference, check: https://doi.org/10.1073/pnas.2011828117); Line 131: Again, for the analysis of oxidative stress it would be important to investigate a tissue specific expression pattern and measurement of ROS markers. Investigating different time points during the exposure also adds to the mechanistic understanding. Do all tissues respond in the same way? In which tissue does an increase in ROS generation start? How? Does it spread to other tissues? By which mechanisms is it generated; Results in general: Tissue specific analyses and more time points can provide a better understanding of how sublethal imidacloprid doses impact growth and survival. Thinking about the doses of choice in light of what bees might be exposed is also important. The mechanistic understanding is missing in the paper, and without it the study does not add much in comparison to previous ones.

Thank you very much for your valuable comments. As you pointed out, the intensity of P450 detoxification and oxidative stress varies considerably between tissues. When checking whole larva gene expression, the tissue-specific profiles become diluted, which is detrimental to elucidating mechanisms. In this study, we encountered technical barriers in obtaining independent samples of specific tissues for anatomical sampling. As a result, we had to forego analysis of some specific tissues, including the tissue-differentiated analyses of P450 gene expression patterns and ROS markers that you mentioned. We only examined larval overall detoxification and antioxidant responses to imidacloprid toxicity. While we do not believe that data from specific tissues are fully representative of the complex overall picture of larvae, there is no doubt that the decision to study larvae as a whole does not contribute to our complete understanding of the mechanisms by which imidacloprid causes larval developmental retardation and larval responses to imidacloprid toxicity. In addition, the fact that this study only analyzed one-time points during imidacloprid exposure and did not design and comparatively analyze different time points limits our complete understanding of the above mechanisms. In summary, as you have pointed out, tissue-specific analyses and more time points could better understand how sublethal doses of imidacloprid affect growth and survival. In future studies, we will overcome the technical challenges and refer to your suggestions for further systematic and in-depth mechanistic studies specifically targeting imidacloprid toxicity in different tissues at different exposure times and incorporate your suggestions, such as whether the response is consistent across all tissues, the origin of the increase in ROS production, how it increases, whether it spreads to other tissues, and the underlying mechanisms into the next experimental design. Again, Thank you for your constructive and valuable comments, which have provided valuable insight for our study on mechanisms. Undoubtedly, these comments will enhance the innovativeness of our study and greatly facilitate our future research.

Q10: Line 236: The conclusion that mitochondrial dysfunction is taking place is not well corroborated. Are there changes in mitochondrial aconitase activity to suggest the mitochondrial origin of ROS? How do mitochondria look like under electron microscopy? Evidence for mitochondrial damage from functional assays? Could the ATP reduced levels be caused by increased consumption by other systems, instead of reduced production? Without functional assays to demonstrate mitochondrial dysfunction the indirect measurements of gene expression at most suggest expression perturbations in mitochondria for the point in time when gene profiles were examined.

Thank you for the comments. Based on the data of the present study, i.e., suppression of mitochondrial oxidative phosphorylation (COX17, NDUFB7) and expression of genes of its alternative glycolytic pathways (Gapdh, Oscillin), as well as a decrease in the ATP content, suggests that imidacloprid exposure leads to impaired energy metabolism in larvae and not to mitochondrial dysfunction. We have corrected this uncritical language presentation error. Please see the lines 267 and 275 red text in our revised version. We hope that this correction will meet with your approval.

Q11: Though not the aim of the study, an important step forward would be to investigate whether these doses that do not impact survival but cause growth retardation could affect the many stereotypical behaviors displayed by the worker bees when they reach the adult life. Without this sort of analysis, it is difficult to stablish whether the doses tested will impact the colony health.

Thank you very much for your valuable suggestions, which give us broader ideas for our subsequent, more in-depth work on the mechanism of toxicity. Inspired by your suggestion, we plan to conduct further studies to investigate the effects of different levels of imidacloprid exposure on the developmental process of bee larvae and the underlying mechanism of toxicity. We will also investigate the intrinsic link between this juvenile toxicity and behavioral and physiological defects in adult individuals.

Q12: Line 376: the authors do not provide a link to their hypothesis that increased P450, and antioxidant response is reducing larvae nutrient supply.

Thank you for your comment. I apologize for not fully understanding your point. If you mean that the hypothesis proposed in this study that increased P450 and antioxidant responses reduce larval nutrient energy supply is not well-founded, we have already addressed this in the previous paragraph. See Figure 7 and lines 395-399 for more details in our revised manuscript.

Q13: Line 393: Were the colonies single-cohort? Were the frames from different hives mixed together to create the experimental groups? Or each experimental group comes from a different frame/colony? This information is important to establish how much genetic variation might exist between the different experimental groups.

Thank you for your comment. In this study, the selected colonies were healthy and not exposed to pathogens or pesticides. Two-day-old larvae from the same frames of the same hive were individually transferred to sterile 24-well cell culture plates. The plates contained a standard diet containing royal jelly, glucose, fructose, water, and yeast extract. We have included the above text in our revised manuscript. Please see the lines 430-432 red text in the revised manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Imidacloprid is a global health threat that severely poisons the economically and ecologically important honeybee pollinators, Apis mellifera. However, its effects on developing bee larvae remain largely unknown. Our pilot study found that imidacloprid causes developmental delay in bee larvae, but its toxicological mechanisms are not yet fully understood. In this study, we exposed bee larvae to imidacloprid at environmentally relevant concentrations of 0.7, 1.2, 3.1, and 377 ppb. A significant dose-dependent delay in the larval development was observed, including reductions in body mass, width, and growth index; however, imidacloprid did not affect larval survival and food consumption. Predominant toxicological effects induced by high concentrations of imidacloprid (377 ppb) included disruption of neural transmission, induction of oxidative stress and gut structural damage and apoptosis, inhibition of developmental regulatory hormones and genes, disruption of proteolysis, amino acid transport, protein synthesis, carbohydrate catabolism, oxidative phosphorylation, and glycolysis energy production. In parallel, we observed that the larvae employed antioxidant defense and P450 detoxification mechanisms to counteract imidacloprid. Ultimately, this study provides the first evidence that environmental exposure to imidacloprid can retard bee larval growth and development by disrupting molting regulation and limiting the metabolism and utilization of dietary nutrients and energy. These findings have broader reference implications for studies assessing pesticide hazards to other juvenile animals.

Revised version: Imidacloprid is a global health threat that severely poisons the economically and ecologically important honeybee pollinator, Apis mellifera. However, its effects on developing bee larvae remain largely unexplored. Our pilot study showed that imidacloprid causes developmental delay in bee larvae, but the underlying toxicological mechanisms remain incompletely understood. In this study, we exposed bee larvae to imidacloprid at environmentally relevant concentrations of 0.7, 1.2, 3.1, and 377 ppb. There was a marked dose-dependent delay in larval development, characterized by reductions in body mass, width, and growth index. However, imidacloprid did not affect on larval survival and food consumption. The primary toxicological effects induced by elevated concentrations of imidacloprid (377 ppb) included inhibition of neural transmission gene expression, induction of oxidative stress, gut structural damage, and apoptosis, inhibition of developmental regulatory hormones and genes, suppression of gene expression levels involved in proteolysis, amino acid transport, protein synthesis, carbohydrate catabolism, oxidative phosphorylation, and glycolysis energy production. In addition, we found that the larvae may use antioxidant defenses and P450 detoxification mechanisms to mitigate the effects of imidacloprid. Ultimately, this study provides the first evidence that environmentally exposed imidacloprid can affect the growth and development of bee larvae by disrupting molting regulation and limiting the metabolism and utilization of dietary nutrients and energy. These findings have broader implications for studies assessing pesticide hazards in other juvenile animals.

Introduction

Earlier version: Hazardous pesticides are extensively used worldwide to safeguard crop production, particularly in developing countries that depend on agriculture. However, exposure to these chemicals has become a global health concern. Imidacloprid, introduced in 1991, is known for its high animal toxicity and has received international attention (Jeschke et al., 2011). The World Health Organization and the U.S. Environmental Protection Agency have listed it as an acute oral and dermal toxicant for animals. Overuse of imidacloprid can lead to environmental contamination, endangering human and animal health through food chains, poisoning non-target pollinators, and disrupting ecological balances (Tudi et al., 2021). Adults exposed to high doses of imidacloprid experience dizziness, apathy, movement disorders, labored breathing, and temporary growth retardation. Imidacloprid also causes adverse congenital, hematological, hepatic, and renal effects, along with degenerative changes in various organs (Khan et al., 2010; Shaw et al., 2014; Yang et al., 2014). In rats and rabbits, imidacloprid causes severe liver, thyroid, and body weight problems, reproductive toxicity, developmental retardation, and neurobehavioral deficits, even at low doses(Anatra-Cordone M, 2005). Imidacloprid also affects the singing ability of birds, making it difficult for them to attract mates and reproduce. Imidacloprid represents a significant contributing factor to the decline in the number of terrestrial insects since 1925, which has declined by an average of 9% per decade (Cardoso et al., 2020; Sánchez-Bayo and Wyckhuys, 2019).

Revised version: Hazardous pesticides are used extensively around the world to protect crop production, particularly in agriculture-dependent developing countries. However, exposure to these chemicals has become a global health concern. Imidacloprid, introduced in 1991, is known for its high animal toxicity and has received international attention (Jeschke et al., 2011). The World Health Organization and the U.S. Environmental Protection Agency have listed it as an acute oral and dermal toxicant for animals. The excessive use of imidacloprid can lead to environmental contamination, endangering human and animal health through food chains, poisoning non-target pollinators, and disrupting ecological balances (Tudi et al., 2021).

Added in the revised version: Imidacloprid is also known for its adverse effects on human health. Adults exposed to high doses of imidacloprid (9.6%) experience dizziness, apathy, movement disorders, labored breathing, and temporary growth retardation (Wu et al., 2001). Imidacloprid also causes adverse congenital, hematological, hepatic, and renal effects, along with degenerative changes in various organs (Khan et al., 2010; Shaw et al., 2014; Yang et al., 2014). Imidacloprid even at low doses causes severe liver, thyroid, and body weight problems, reproductive toxicity, developmental retardation, and neurobehavioral deficits in rats and rabbits (Anatra-Cordone M, 2005). Additionally, imidacloprid affects the singing ability of birds, making it difficult for them to attract mates and reproduce. It is a significant contributor to the decline in terrestrial insect populations since 1991, averaging 9% per decade (Cardoso et al., 2020; Sánchez-Bayo and Wyckhuys, 2019). Honey bees (Apis mellifera L.) are valuable model organisms in insects with significant economic and ecological importance as crop pollinators and in maintaining ecological balance (Zheng and Fu-Liang, 2009). However, the widespread use of pesticides, including imidacloprid, has resulted in colony collapse disorder (CCD), causing severe impacts on the beekeeping industry and ecological balance (Mahé et al., 2021). Studies have shown that imidacloprid is applied as a seed coating (Simon-Delso et al., 2015) and accumulates in the plant (Dively and Kamel, 2012; Rocha et al., 2022), exposing non-target bees to imidacloprid residues during pollen and nectar collection (Tong et al., 2018). Even low levels of systemic imidacloprid cause sublethal effects in adult bees, including neurotoxicity (Belzunces et al., 2012), behavioural changes (Karahan et al., 2015; Medrzycki et al., 2003), reproductive impairment (Chaimanee et al., 2016), reduced immunity (Pettis et al., 2013), apoptosis and autophagy (Carneiro et al., 2022), and shortened lifespan (Anderson and Harmon-Threatt, 2021). A recent finding shows that imidacloprid induces oxidative stress in honey bees (Balieira et al., 2018). This oxidative stress is a result of the long-term synergistic effect between imidacloprid and its target receptors (nAChRs) (Makoto et al., 2020) in bees, which leads to an excessive influx of Ca2+ into the bee brain (Tahirare, 2013), disrupting Ca2+ homeostasis and ultimately causing mitochondrial dysfunction and increased mitochondrial ROS generation, exacerbating oxidative stress.

Added in the revised version: The development and health status of juvenile animals play a critical role in population structure and future growth (Gill et al., 2012). However, they are more susceptible to the toxic effects of pesticides than adults. Despite this, there is limited information on the toxic effects of imidacloprid on juvenile animals. Available reports mainly focus on the epistatic effects of toxicity, with limited investigation into the involved biochemical mechanisms. Imidacloprid is known to pose a potential health risk to children through dietary exposure (Anderson and Harmon-Threatt, 2021), and may cause lead to developmental abnormalities in amphibians (Samojeden et al., 2022) and fish (Islam et al., 2019). Furthermore, negative effects on pupal development in insects have been observed (Jia and Li, 2023). For instance, cat fleas (Ctenocephalides felis) larvae treated with imidacloprid become immobile, and their intestines show pulsatile movements for up to one hour, eventually dying after two hours (Mehlhorn et al., 1999). Additionally, imidacloprid exposure can affect the average weight gain per day in Harmonia axyridis (Coleoptera: Coccinellidae) larvae (Charles et al., 2000). It has also been found that imidacloprid exposure affects various cellular aspects, such as mitochondria, energy, lipids, and the transcriptome in Drosophila (Felipe et al., 2020), and prolongs the larval pupation duration and affects forewing development in the butterfly Pieris brassicae (Whitehorn et al., 2018).

Added in the revised version: In bees, larvae play a crucial role in the growth and development of colonies. However, the broader toxic effects of imidacloprid on bee larvae remain to be investigated. Limited reports suggest that developing bee larvae are exposed to imidacloprid through contact with contaminated hives or ingestion of contaminated food (Bhme et al., 2018). However, their survival, developmental rate, and body weight are not affected (Dai et al., 2019). Nevertheless, other studies have shown that imidacloprid can lead to impaired olfactory learning behavior (Peng and Yang, 2016), and other behavioral abnormalities (Wu et al., 2017). The current understanding is that honey bee larvae are more tolerant to imidacloprid than adults, as they can survive treatment and develop into adults (Dai et al., 2017). However, they may experience developmental delays, such as delayed plumage (Woyciechowski and Moroń, 2009; Wu et al., 2011), and reduced rates of capping, pupation, and eclosion (Yang et al., 2012). A recent study using next-generation sequencing has indicated that sublethal imidacloprid treatment during the larval stage causes changes in gene expression in larvae, pupae, and adults, suggesting a persistent sublethal impact on honey bee development (Chen et al., 2021). Furthermore, a latest study has demonstrated that imidacloprid induces differential gene expression in bumble bee larvae, resulting in unique expression patterns related to various biological processes such as starvation response, cuticle genes, neural development, and cell growth (Martín-Blázquez et al., 2023).

Added in the revised version: Although progress has been made in understanding the toxic effects of imidacloprid on honey bee larvae, the molecular basis of this toxicity remains poorly understood. Further research is needed to uncover the mechanisms and pathways involved in the developmental toxicity of imidacloprid in honey bee larvae. In our pilot study, we observed a significant developmental delay in bee larvae exposed to imidacloprid; However, the molecular and biochemical mechanisms underlying this delay have not been systematically investigated. Therefore, this study aims to comprehensively evaluate the toxicity of imidacloprid on honey bee larvae at multiple levels, including histology, neurotoxicity, molting regulation, oxidative stress, detoxification, nutrition, and energy metabolism. Our study provides first-hand evidence of the developmental delay caused by imidacloprid in bee larvae. Because these mechanisms involved are relatively conserved in animals, our findings have broader implications for understanding and assessing the potential risk of continuous exposure to imidacloprid during animal development.

Removed in the revised version: The development and health status of juvenile animals play a critical role in population structure and future growth (Gill et al., 2012), but they are more susceptible to the toxic effects of pesticides than adults. However, there is a paucity of information on the assessment of toxic effects of imidacloprid on juvenile animals. Available reports primarily describe the epistatic traits influenced by toxicity, with limited investigation into the biochemical mechanisms involved. Studies have shown that imidacloprid poses a potential health risk to children through dietary exposure (Anderson and Harmon Threatt, 2021) and can lead to developmental abnormalities in amphibians (Samojeden et al., 2022) and fish (Islam et al., 2019) as well as negative effects on insect pupal development (Jia and Li, 2023). In insects, the honeybees (Apis mellifera L.) are valuable model organisms with significant economic and ecological importance as pollinators of crops and maintenance of ecological balance (Zheng and Fu-Liang, 2009). However, the widespread use of pesticides, including imidacloprid, has resulted in colony collapse disorder (CCD), causing serious impacts on beekeeping industry and ecological balance (Mahé et al., 2021). Studies have shown that Imidacloprid is applied as seed coating (Simon-Delso et al., 2015) and accumulate in the plant (Dively and Kamel, 2012; Rocha et al., 2022), exposing non-target bees to imidacloprid residues during pollen and nectar collection (Tong et al., 2018). Even low levels of systemic imidacloprid cause sublethal effects in adult bees, including neurotoxicity (Belzunces et al., 2012), behavioral changes (Karahan et al., 2015; Medrzycki et al., 2003), reproductive impairment (Chaimanee et al., 2016), immunity (Pettis et al., 2013), apoptosis and autophagy (Carneiro et al., 2022), and shortened lifespan (Anderson and Harmon-Threatt, 2021). Larvae is the future of colony growth and development, yet the broader toxic effects of imidacloprid on larvae remain unknown. Limited reports suggest that developing bee larvae are exposed to imidacloprid through contact with contaminated hives or ingesting contaminated food (Bhme et al., 2018), leading to delayed development, reduced rates of capping, pupation, and eclosion (Yang et al., 2012), and impaired olfactory learning behavior (Peng and Yang, 2016). In the pilot study, we found that imidacloprid exposure caused obvious developmental delay in bee larvae, but the molecular and biochemistry mechanisms remain to systematically investigated. To this end, this study comprehensively evaluated the toxicity of imidacloprid on honey bee larvae at multiple levels including histology, neurotoxicity, molting regulation, oxidative stress, detoxification, nutrition, and energy metabolism, and provides first-hand evidence for imidacloprid causing developmental delay in bee larvae. Because the mechanisms involved are relatively conserved in animals, our findings have broader implications for understanding and assessing the potential risk of continuous exposure to imidacloprid during animal development.

Results

Imidacloprid causes bee larval developmental retardation

Earlier version: During the 96-hour observation period, larval survival in the four imidacloprid-exposed groups was not significantly different from that in the control group. However, with increasing exposure concentration and duration, larval developmental progress, weight and width were significantly reduced (Fig. 1A-1E). Further analysis focused on the most adverse exposure level of 377 ppb, which did not affect survival, feeding (Fig. 1C, 1F, and 1G), and the time taken for larvae to successfully enter the pupal stage (Fig. 1H), but resulted in slow developmental obstacles and lower initial rate of successful pupation and growth index. In addition, imidacloprid also caused consistently lower weight (Fig. 1I) and width (Fig. 1J) of larvae up to the pupal stage (Fig. 1B).

Revised version: The 96-hour observation period showed no significant difference in larval survival between the four groups exposed to imidacloprid and the control group. However, as the exposure concentration and duration increased, here was a significant decrease in larval developmental progress, weight, and width (Fig. 1A-1E). Further analysis was conducted on the most adverse exposure level of 377 ppb, which is the highest reported environmental residue concentration detected in bee products, such as beeswax (Kapoor et al., 2014). It is important to note that while the tested doses of the 377 ppb was found in bee products, it is likely that the bees themselves were exposed to even higher doses. Out result demonstrate that exposure to 377 ppb imidacloprid did not have a significant impact on survival, feeding (Fig. 1C, 1F, and 1G), and the time required for larvae to enter the pupal stage successfully (Fig. 1H). However, it did lead to slow developmental progress, a lower initial rate of successful pupation, and lower growth index. Additionally, imidacloprid consistently resulted in lower larval weight (Fig. 1I) and width (Fig. 1J) up to the pupal stage (Fig. 1B).

Imidacloprid causes larval developmental retardation in Apis mellifera. Larvae were orally administered imidacloprid (IMI) or artificial food (CK) beginning at three-day-old, followed by developmental monitoring for 96 hours. (A) Larval developmental phenotypes at 0.7, 1.2, 3.1, and 377 ppb environmental concentrations of imidacloprid exposure; (B) Larval developmental phenotype at 377 ppb imidacloprid exposure; (C) Effect of 0.7, 1.2, 3.1, and 377 ppb imidacloprid on larval survival; (D) and (E) body width and weight of larvae exposed to 0.7, 1.2, 3.1, and 377 ppb imidacloprid for 96 hours; (F) and (G) statistics for daily (F) and three-day cumulative feeding (G) for larvae exposed to 377 ppb imidacloprid; (H) larval development time from three-day-old to pupal stage; (I) The number of larvae that successfully reached the pupal stage as a percentage of the total initial sample size; (J) The growth index of larvae. Statistical significance was set at *P < 0.05 and **P < 0.01.

Imdacloprid neurotoxicity in bee larvae

Earlier version: Under imidacloprid environmental stress at 377 ppb, AChE activity, which is closely related to the disruption of neurotransmission, was significantly inhibited in larvae (Fig. 2A). Neither acetylcholine receptor alpha 1 (Alph1) nor acetylcholinesterase 2 (Ace2) gene expression was significantly different from the control. In contrast, acetylcholine receptor alpha 2 (Alph2) and acetylcholinesterase 1 (Ace1) gene expression was significantly higher than that in the control group (Fig. 2B).

Revised version: Under imidacloprid environmental stress at 377 ppb, AChE activity, which is closely related to the disruption of neurotransmission, was significantly inhibited in larvae (Fig. 2A). Neither acetylcholine receptor (nAChR) alpha 1 (Alph1) nor acetylcholinesterase 2 (Ace2) gene expression was significantly different from the control. In contrast, acetylcholine receptor alpha 2 (Alph2) and acetylcholinesterase 1 (Ace1) gene expression was significantly higher than that in the control group (Fig. 2B).

Imidacloprid neurotoxicity in Apis mellifera larvae. (A) AChE activity and (B) nerve conduction-related gene expression analysis in larvae exposed to 377 ppb imidacloprid for 72 hours or a control group. IMI is the larvae exposed to imidacloprid for 72 hours. CK is the artificial food control group. Statistical significance was set at *P < 0.05 and **P < 0.01.

Imidacloprid disrupts the homeostasis of developmental regulation in larvae

In imidacloprid-exposed larvae, the developmental regulatory gene juvenile hormone acid methyl transferase (JHAMT) and the transcription factor broad complex (Br-c) were downregulated, while vitellogenin (Vg) gene expression was upregulated by seven-fold in comparison to control groups (Fig. 3A). In parallel, imidacloprid exposure decreased the major regulator of insect metamorphosis 20-hydroxyecdysone (20E) titers but had no effect on preventing premature metamorphosis juvenile hormone (JH-3) titers (Fig. 3B).

Effects of imidacloprid on the homeostasis of developmental regulation in Apis mellifera larvae. (A) Relative gene expression and (B) hormone levels of developmental regulatory-related genes in larvae exposed to 377 ppb imidacloprid for 72 h or a control group. Statistical significance was set at *P < 0.05 and **P < 0.01.

Imidacloprid induces detoxification in larvae

According to the qRT-PCR analysis, the expression levels of cytochrome P450 monooxygenase family members such as CYPq1, CYPq2, CYP450, CYP6AS14, CYP4G11, and CYP306A1, except for CYPq3, were observed to be significantly upregulated by several fold in imidacloprid-induced developmental retardation larvae. Among them, CYPq1 showed an abnormal 94-fold increase (Fig. 4).

Imidacloprid induced a detoxification response in Apis mellifera larvae. Relative expression levels of cytochrome P450 monooxygenase family genes in larvae exposed to imidacloprid for 72 hours. IMI is the larvae exposed to imidacloprid for 72 hours. CK is the artificial food control group. Statistical significance was set at *P < 0.05, **P < 0.01, and ***P < 0.001.

Imidacloprid toxicity causes oxidative stress and induces antioxidant defense in larvae

Larvae exposed to imidacloprid exhibited a significant increase in reactive oxygen species (ROS) levels (Fig. 5A) and upregulation of the antioxidant genes GPX and Trx expression (Fig. 5C), as well as increased activities of the antioxidants CAT, SOD and GSH (Fig. 5B), resulting in a significant increase in total antioxidant capacity (T-AOC) (Fig. 5B) compared to the control group. Under these conditions, MDA levels, used as a marker of lipid peroxidation, were observed to increase, while protein carbonylation damage indicator PCO levels remained unchanged (Fig. 5A).

Imidacloprid exposure causes oxidative stress, gut apoptosis, and tissue structural damage while inducing antioxidant defenses in Apis mellifera larvae. (a) Analysis of oxidative stress and damage in larvae exposed to imidacloprid for 72 hours. ROS levels indicate the degree of oxidative stress. MDA is a marker for lipid peroxidation. PCO levels reflect the extent of protein carbonylation damage; (b) Activity of antioxidants in larvae exposed to imidacloprid for 72 hours; (c) The relative expression levels of antioxidant genes in larvae exposed to imidacloprid for 72 hours; (d) Histological sections of the larval gut stained with H&E. The circled inserts in the figure show magnified views of the muscle layer. The blue arrow indicates the cell nucleus of the muscle layer. The boxed inset in the figure is a magnified view of the basal layer (black arrow), peritrophic membrane (red arrow), and food residues (green arrow). IMI is the larvae exposed to imidacloprid for 72 hours. CK is the artificial food control group. Statistical significance was set at *P < 0.05, **P < 0.01, and ***P < 0.001.

Imidacloprid induces gut apoptosis and tissue damage in larvae

Hematoxylin-eosin staining revealed a significant reduction in the number of cells in the muscular layer of the larval gut exposed to imidacloprid with clear apoptotic characteristics compared to the standard feeding control groups (Fig. 5D). This was evidenced by increased nuclear staining and compacted chromatin, as well as disrupted cell arrangement in the basal layer of the gut, increased nuclear staining of basal lamina cells, and apoptotic signals. The integrity of the peritrophic membrane structure of the gut was better in the control group, while imidacloprid-exposed larvae were still in the formation stage, resulting in the accumulation of undigested food residues in their guts (Fig. 5D).

Imidacloprid impairs larval digestion and breakdown of dietary nutrients

Exposure to imidacloprid did not significantly affect the daily food intake of developmentally delayed larvae compared to controls (Fig. 1F and 1G). However, imidacloprid toxicity suppressed the expression of nutrient catabolism genes, including alpha-amylase (α-Am) and alpha-glucosidase (α-Glu), which are involved in carbohydrate catabolism (Fig. 6A), carboxypeptidase (CPs) and aminopeptidase (APs) involved in proteolysis (Fig. 6B), and glycine tRNA ligase (Aats-Gly) and tyrosine tRNA ligase (Aats-Tyr) involved in amino acid transport (Fig. 6C).

Imidacloprid toxicity inhibited the expression of genes involved in nutrient catabolism, protein synthesis, and energy metabolism and reserves. (A) Carbohydrate catabolism; (B) Proteolysis; (C) Amino acid transport; (D) Mitochondrial oxidative phosphorylation; (E) Glycolysis; (F) Protein synthesis. (G) The total contents of ATP, glycogen (H), and protein (I). IMI is the larvae exposed to imidacloprid for 72 hours. CK is the artificial food control group. Statistical significance was set at *P < 0.05, **P < 0.01, and ***P < 0.001.

Imidacloprid causes restriction of larval protein synthesis and energy metabolic dysfunction

RT-qPCR analysis revealed that imidacloprid exposure reduced the expression of critical genes involved in energy metabolism, including cytochrome c oxidase copper chaperone (COX17) and NADH dehydrogenase [ubiquinone] 1β subcomplex subunit 7 (NDUFB7), in developmentally delayed larvae (Fig. 6D). Imidacloprid exposure also inhibited the expression of genes related to glycolysis and energy production, such as glyceraldehyde-3-phosphate dehydrogenase (Gapdh) and glucosamine-6-phosphate isomerase (Oscillin) (Fig. 6E). In addition, the expression of transcription initiation factors related to protein synthesis, including translation initiation factor 3 subunit M (Tango7), translation initiation factor 3 subunit B (elF3-S9), and translation initiation factor 3 subunit A (elF3-S10), was significantly suppressed by imidacloprid toxicity (Fig. 6F). Finally, the total levels of ATP, glycogen and protein, which are closely related to energy production and reserves, were significantly lower in developmentally delayed larvae than in controls (Fig. 6G-6I).

Correlation between developmental phenotypic changes and molecular features caused by imidacloprid toxicity

To better understand the mechanism of imidacloprid-induced larval developmental retardation, we performed a comprehensive correlation analysis between epigenetic growth traits and molecular characteristics (Fig. 7). The results showed that in larvae with imidacloprid-induced developmental retardation, body weight, width, growth index, and developmental rate were positively correlated with energy metabolism and reserves, nutrient catabolism, protein synthesis, and developmental regulation. Conversely, negative correlations were observed between these growth traits and parameters associated with detoxification and antioxidant defense processes that require additional energy expenditure. Moreover, oxidative stress and damage, detoxification, and antioxidant defense all exhibited significant positive correlations (Fig. 7).

Correlation between phenotypes and molecular characteristics in imidacloprid-exposed larvae with developmental retardation. Data were normalized using GraphPad software, followed by a Pearson correlation coefficient calculation using SPSS, and a heatmap was generated.

Discussion

Toxic effects and the molecular basis of environmental concentrations of imidacloprid on bee larvae

Previous research has shown that bee larvae are more susceptible to the adverse effects of pesticides than adult bees (Tomé et al., 2020). In this study, larvae were exposed to environmental concentrations of imidacloprid, resulting in sublethal effects such as decreased body weight, width, developmental rate, and growth index (Fig. 1). Our results support the notion that pesticides can affect larval development (Dai et al., 2018; Zhu et al., 2014). However, current research has primarily focused on characterizing the effects of imidacloprid on larval phenotype and behavior, and the molecular mechanisms underlying its ecotoxicological effects remain poorly understood.

To obtain molecular evidence of imidacloprid toxicity to bee larvae, we first examined the expression of acetylcholine receptors and AChE activity. The results showed that imidacloprid exposure increased the expression of the Alph2 gene and inhibited AChE activity (Fig. 2). Normally, acetylcholine stimulates nerve impulses by binding to the Alph receptor, and AChE degrades acetylcholine to terminate nerve conduction. However, imidacloprid competitively binds to AChE receptors, inhibits AChE activity, and causes continuous excitation of nerve conduction, ultimately leading to paralysis and death (Katić et al., 2021; Shan et al., 2020). Therefore, our results strongly suggest that imidacloprid is neurotoxic to bee larvae.

The growth and development of juvenile animals depends on the digestion, absorption and utilization of food nutrients. Therefore, we next investigated whether imidacloprid affects the digestion and utilization of dietary by larvae. The results indicate that exposure to imidacloprid suppressed the hydrolysis genes CPs and APs and the amino acid transport pathway genes Aats-Gly and Aats-Tyr in larvae (Fig. 6B and 6C). Since these genes are responsible for cleaving dietary proteins into absorbable peptides and amino acids and transporting them for metabolism (2019), inhibiting the expression of these genes would severely affect larval digestion, resulting in protein deficiency and impaired metabolic activity. In addition, imidacloprid also affected the expression of α- Am and α-Glu genes (Fig. 6A), which are involved in carbohydrate digestion, indicating that imidacloprid also affects the digestion of dietary carbohydrates. Indeed, histological observations also confirmed that a large amount of undigested food remained in the intestine of larvae exposed to imidacloprid (Fig. 5D). In conclusion, imidacloprid inhibits the digestion and utilization of dietary proteins and carbohydrates by larvae, which is extremely detrimental to their growth and development.

Earlier version: In addition to nutrition, energy metabolism is critical for the maintenance of all life activities. In this study, imidacloprid exposure inhibited the expression of genes related to mitochondrial oxidative phosphorylation (COX17, NDUFB7) and its alternative glycolytic pathway (Gapdh, Oscillin). Mitochondrial oxidative phosphorylation is the central pathway for ATP generation during energy metabolism (Tian et al., 2020a), while glycolysis is an alternative mechanism for ATP synthesis (Jiang et al., 2020). Since both biochemical processes are essential for energy production, their inhibition resulted in reduced ATP levels in bee larvae (Fig. 6G). These results strongly suggest that imidacloprid induces mitochondrial dysfunction leading to insufficient energy production in larvae. Since larval energy production is insufficient, this is likely to lead to a depletion of energy reserves. Indeed, further analysis revealed a significant decrease in glycogen and protein energy reserves in larvae exposed to imidacloprid (Fig. 6H and 6I). These findings are consistent with previous reports that exposure to pesticides such as endosulfan, chlorpyrifos, insecticides, and fipronil can lead to a reduction in energy reserves in organisms (Bouayad et al., 2012a; Dutra et al., 2009; Radwan et al., 2008; Rambabu and Rao, 1994; Ribeiro et al., 2001). In conclusion, our results suggest that imidacloprid induces mitochondrial dysfunction and affects energy production and reserves in bee larvae.

Revised version: In addition to nutrition, energy metabolism is critical for the maintenance of all life activities. In this study, imidacloprid exposure inhibited the expression of genes related to mitochondrial oxidative phosphorylation (COX17, NDUFB7) and its alternative glycolytic pathway (Gapdh, Oscillin). Mitochondrial oxidative phosphorylation is the central pathway for ATP generation during energy metabolism (Tian et al., 2020a), while glycolysis is an alternative mechanism for ATP synthesis (Jiang et al., 2020). Since both biochemical processes are essential for energy production, their inhibition resulted in reduced ATP levels in bee larvae (Fig. 6G). These results strongly suggest that imidacloprid induces energy metabolism dysfunction leading to insufficient energy production in larvae. Since larval energy production is insufficient, this is likely to lead to a depletion of energy reserves. Indeed, further analysis revealed a significant decrease in glycogen and protein energy reserves in larvae exposed to imidacloprid (Fig. 6H and 6I). These findings are consistent with previous reports that exposure to pesticides such as endosulfan, chlorpyrifos, insecticides, and fipronil can lead to a reduction in energy reserves in organisms (Bouayad et al., 2012a; Dutra et al., 2009; Radwan et al., 2008; Rambabu and Rao, 1994; Ribeiro et al., 2001). In conclusion, our results suggest that imidacloprid induces energy metabolism dysfunction and affects energy production and reserves in bee larvae.

Insects rely primarily on cytochrome P450 detoxification to resist toxic compounds (Li et al., 2007), and the level of CYP450 expression indicates the severity of pesticide toxicity. While adult worker bees have shown a positive response to imidacloprid (Mao et al., 2009), it is uncertain whether bee larvae can respond to pesticide toxicity. In this study, we found that seven cytochrome P450 family transcripts, including CYPq1, CYPq2, CYP450, CYP6AS14, CYP4G11, and CYP306A1, were significantly upregulated in bee larvae, with CYP9q1 being 94-fold higher than the control group (Fig. 4). These results are consistent with previous findings that the CYP450 family genes CYP6AS3, CYP6AS4, and CYP9S1 were significantly upregulated in adult worker bees exposed to quercetin, coumaphos, and fluvalinate (Mao et al., 2011) and that CYP9Q1, CYP9Q2, and CYP9Q3 help bees detoxify fluvalinate and coumaphos (Gregorc et al., 2018). This suggests that the P450 detoxification system is already established in bee larvae and they respond positively by activating the detoxification function of CYP450 upon exposure to imidacloprid.

Oxidative stress is the excessive production of ROS in aerobic organisms (Felton, 1995; Ozcan and Ogun, 2015). Pesticides, including coumaphos (Gregorc et al., 2018), fipronil (Paris et al., 2017), paraquat (Li-Byarlay et al., 2016), organophosphorus, pyrethroid, organochlorinated (Chakrabarti et al., 2015), and chlorpyrifos (Shafiq-ur-Rehman et al., 2012), have been shown to cause oxidative stress in bees. In this study, exposure to imidacloprid resulted in increased levels of ROS and MDA in larvae (Fig. 5A), indicating severe oxidative stress and lipid damage. Organisms have evolved antioxidant defense mechanisms to scavenge excessive ROS and resist oxidative stress, including increasing levels of antioxidants such as CAT, SOD, GSH, TrxR, and GSH (Dong et al., 2013; Tian et al., 2020b). In this regard, we have recently confirmed this in worker bees (Li et al., 2022), which is also well reflected in the present study. We observed a significant increase in the activities of antioxidant enzymes SOD and CAT, gene expression of Trx and Gpx, and GSH levels in larvae after exposure to imidacloprid (Fig. 5B and 5C). These results show that the antioxidant system in developing larvae is established and used to resist oxidative stress caused by imidacloprid, similar to that in adult worker bees. Notably, compared to our recent study in adult bees (Li et al., 2022), the larvae in this study exhibited more robust antioxidant activity when exposed to higher concentrations of imidacloprid toxicity than worker bees. However, the underlying mechanism requires further investigation. The underlying mechanisms involved require further study.

Notably, in addition to routine antioxidants, imidacloprid induces a 7-fold increase in Vg expression (Fig. 3A). Vg is known to provide nutrition to developing embryos (Jorgensen et al., 2009) and to help animals cope with pollution-induced stress (Matozzo et al., 2008). In bees, Vg plays an important antioxidant role by protecting other hemolymph molecules from ROS and increasing cellular oxidative stress tolerance (Corona et al., 2007). Previous studies have shown a correlation between high Vg expression in adult worker bees and resistance to oxidative stress resulting from exposure to the pesticide paraquat (Seehuus et al., 2006). In this study, our results support this observation and suggest that larvae may also resist oxidative stress caused by imidacloprid through increased Vg production.

Mechanisms of imidacloprid-induced developmental retardation in bee larvae

Impaired juvenile development can have a significant impact on the growth and sustainability of entire populations. Deviations in juvenile development due to environmental stress have recently received considerable attention. Drosophila melanogaster exposed to heat stress exhibit developmental delays due to the failure of early embryos to fully inhibit the synthesis of non-heat shock proteins (Bergh and Arking, 1984). Similarly, exposure to nutrient and heavy metal stress can cause delayed or arrested development in Caenorhabditis elegans due to altered gene expression and disruption of developmental signaling pathways (Carranza-García and Navarro, 2020; Rashid et al., 2021). In the present study, we also observed typical developmental delays in honeybee larvae exposed to imidacloprid, but the underlying mechanisms remain unclear.

Insect development is regulated by the hormones JH and 20E, which control morphological remodeling events during molting. 20E promotes pupation (Yuan et al., 2020), while JH maintains juvenile characteristics and prevents metamorphosis by antagonizing the effects of 20E (Luo et al., 2021). To elucidate the mechanism underlying imidacloprid-induced larval developmental delay, we first focused on the hormonal regulation of molting development and found that imidacloprid inhibits the expression of JHAMT (Fig. 3A), a positive regulator of JH synthesis (Liu et al., 2018), but did not affect JH titer levels (Fig. 3B). This suggests that the antagonistic effect of JH still restricts the developmental process, thereby maintaining larval morphology. In parallel, imidacloprid toxicity also resulted in reduced ecdysteroid 20E titer and Br-c expression (Fig. 3). The role of 20E is to induce larval molting and pupation by triggering downstream response factors, such as Br-c (Deng et al., 2012). Therefore, our results suggest that larval developmental delay may be caused by imidacloprid reducing 20E titer and inhibiting Br-c expression, thereby blocking molting. Notably, we also observed that a decrease in 20E levels was accompanied by a decrease in AChE activity (Fig. 2A), which is known to cause delayed pupation (He et al., 2012) and developmental arrest (Desneux et al., 2007) in insect larvae. Therefore, our results suggest that imidacloprid neurotoxicity may cause developmental delay in bee larvae by inhibiting the molting hormone 20E. This conclusion is supported by a correlation analysis showing that decreases in 20E titer, Br-c expression and AChE activity were positively correlated with decreased developmental rates and growth index, weight and width growth (Fig. 7). In conclusion, imidacloprid neurotoxicity inhibits the 20E and Br-c genes, thereby blocking molting, which may be an important cause of delayed larval development in honey bees.

Animals employ selective foraging as a survival strategy to avoid ingesting toxic foods that may cause adverse physiological effects (Berenbaum and Johnson, 2015). Although bees cannot actively reject toxins innately, postingestive malaise allows them to learn to avoid ingesting toxic-containing foods with adverse physiological effects (Hurst et al., 2014). Given this, we speculate that it is likely that the honey bee larvae in this study reduced their intake of toxic imidacloprid-containing foods, making it difficult to meet their nutritional and energy needs, which may have resulted in delayed development. However, monitoring of daily food intake did not show a significant difference between larvae exposed to imidacloprid and the control group (Fig. 1F and 1G), indicating that larval food intake was not reduced and larvae did not initiate avoidance behavior toward toxic food. Therefore, food intake was not a critical factor in the delayed larval development caused by imidacloprid.

Nutrition and energy are critical for the growth and development of juvenile animals. Exposure to heavy metals and neonicotinoid pesticides has been shown to affect digestion, energy reserves, production, and utilization in animals (Jiang et al., 2020; Li et al., 2021). Given this, since the lack of the necessary correlation between delayed larval development and food intake in this study prompts us to consider potential problems in larval food and energy utilization that led to delayed development. Consequently, our focus shifted to the investigation of larval nutrition and energy metabolism. Interestingly, we found severe pathological changes in the gut tissues of developmentally delayed larvae exposed to imidacloprid (Fig. 5D). Imidacloprid caused a significant reduction in the tightness of the cell arrangement in the basal layer of the gut, and a drastic decrease in the number of cells. The peritrophic membrane was incompletely formed, and undigested food residues were present in the gut. qRT-PCR analysis revealed a significant decrease in larval digestion and catabolism of dietary proteins and carbohydrates, amino acid transport, oxidative phosphorylation, and glycolysis (Fig. 6D-6F). Biochemical assays showed that imidacloprid stress also decreased ATP levels and total glycogen and protein content, resulting in a lack of energy reserves and increased energy expenditure (Fig. 6G-6I). Correlation analysis indicated a positive relationship between impaired catabolism and dietary nutrient utilization caused by imidacloprid and larval developmental delay (Fig. 7). These evidence strongly suggest that imidacloprid exposure affects the digestion and utilization of nutrients and energy in larvae, resulting in insufficient nutrients and energy for growth and development, which may be another important factor contributing to imidacloprid-induced larval developmental delay.

The Dynamic Energy Budget (DEB) theory suggests that environmental stress can significantly affect the energy balance of organisms, leading to increased energy expenditure (Kooijman, 2009). To cope with stress, animals may deplete their energy reserves, such as glycogen and protein (Bouayad et al., 2012b; Matsukura et al., 2008). In this study, we found that larvae with delayed development caused by imidacloprid had lower energy reserves (ATP, total protein, and total glycogen), and these energy reserves were negatively correlated with larval antioxidant and detoxification (Fig. 7), suggesting that the reduction in energy reserves caused by imidacloprid is related to the energy-consuming activities of CYP450 detoxification and antioxidant defense. To meet the high energy cost demands of environmental stress, animals often have to allocate most of their food intake to defense (Guedes et al., 2006), and this redistribution of energy resources will inevitably lead to an insufficient energy supply for growth and development (Beyers et al., 1999). Therefore, the P450 detoxification and antioxidant defense used by the larvae in this study to resist imidacloprid toxicity will inevitably force them to consume more food resources that should have been used for growth and development, resulting in a severe lack of energy needed for development. This may also be another important reason why imidacloprid causes delayed development in bee larvae.

Added in the revised version: It is worth noting that our investigation into the causes of imidacloprid-induced developmental delay in honey bee larvae has provided limited answers. Future comparative analyses of different tissues and durations of exposure, using genetic manipulation and histochemical methods, would improve our understanding of the effects of sublethal doses of imidacloprid on larval development, growth and survival.

Conclusion

Imidacloprid toxicity caused growth and developmental delay in A. mellifera larvae. The toxic mechanism may include imidacloprid disrupting the regulatory balance of molt development, resulting in restricted development. In addition, the gut damage caused by imidacloprid restricts the metabolism and utilization of food nutrients and energy in the larvae. Third, the additional energy consumed by larval P450 detoxification and antioxidant defenses further reduces the supply of nutrients and energy for growth and development (Fig. 8). This study is the first multi-level investigation of the toxic effects of imidacloprid on A. mellifera larvae. These findings have broader implications for understanding and assessing the threat of harmful pesticides to animal larvae.

A model interpreting the ecotoxicological effects of imidacloprid on Apis mellifera larvae and the mechanisms of imidacloprid-induced developmental retardation. Imidacloprid toxicity caused growth and developmental delay in A. mellifera larvae. It disrupts the regulatory balance of molt development, resulting in restricted development. In addition, the damage caused by imidacloprid restricts the metabolism and utilization of food nutrients and energy in larvae. The additional energy consumed by larval P450 detoxification and antioxidant defenses further reduces the supply of nutrients and energy supply for growth and development. The arrows indicate induction or promotion, the straight lines indicate inhibition, and the green dotted lines indicate speculative inhibition.

Materials and Methods

Larval rearing

Earlier version: The Animal Bioethics Committee at Chongqing Normal University, China approved all experimental protocols. The chosen colonies were healthy and not exposed to pathogens or pesticides. Two-day-old larvae were individually transferred to sterile 24-well cell culture plates filled with a standard diet consisting of royal jelly, glucose, fructose, water, and yeast extract (Aupinel et al., 2005). The larvae were cultured in dark conditions at 35°C and 96% RH, with daily feeding volumes adjusted according to their age (20 μL for three-day-old larvae, 30 μL for four-day-old larvae, 40 μL for five-day-old larvae, and 50 μL for six-day-old larvae). Any remaining food was recorded and removed.

Revised version: The Animal Bioethics Committee at Chongqing Normal University, China approved all experimental protocols. The chosen colonies were healthy and not exposed to pathogens or pesticides. Two-day-old larvae from the same frames of the same hive were individually transferred to sterile 24-well cell culture plates. The plates contained standard food, which included royal jelly, glucose, fructose, water, and yeast extract (Aupinel et al., 2005). The larvae were cultured in dark conditions at 35°C and 96% RH, with daily feeding volumes adjusted according to their age (20 μL for three-day-old larvae, 30 μL for four-day-old larvae, 40 μL for five-day-old larvae, and 50 μL for six-day-old larvae). Any remaining food was recorded and removed.

Experimental design

Earlier version: This study employed four different concentrations of imidacloprid, which were residual concentrations found in honey (0.7 ppb) (Chauzat et al., 2009), bees (1.2 ppb) (Chauzat et al., 2011), pollen (3.1 ppb) (Mullin et al., 2010), and beeswax (377 ppb) (Kapoor et al., 2014). The experiment was conducted with three treatment groups: control group (CK), solvent control group (CKac), and imidacloprid-treated group (IMI), each with three replicates. Each replicate contained 32 bees that were distributed in standard 24-well cell culture plates (1 bee/well). The CK group consisted of larval bees fed acetone-royal jelly glucose-fructose-yeast extract-distilled water (0.00377%-50%-6%-6%-1%-36.99623%, v/w/w/w/w/v). The IMI group consisted of larvae-fed control solution containing imidacloprid (99% active ingredient, Hubei Norna Technology Co. Ltd, Hubei, China) at 0.7 ppb, 1.2 ppb, 3.1 ppb, and 377 ppb. The CKac group served as a solvent control and contained only acetone at a concentration of 0.00377%. A total of 312 larvae were used for survival statistics (96 bees), body weight and width (96 bees), and gene expression, enzyme activity, histochemical content, and histopathologic analysis. Food was administered daily according to age (20 μL for three-day-old larvae, 30 μL for four-day-old larvae, 40 μL for five-day-old larvae, and 50 μL for six-day-old larvae). Dead individuals and food remain were recorded daily, and incubation conditions were 35°C and 96% RH. Surviving larvae were collected 72 h after consumption of the test solution and stored at −80°C for further analysis, including evaluations of gene expression, enzyme activity, and chemical content evaluations. Furthermore, HE-stained sections were used to evaluate any imidacloprid-induced tissue damage and apoptosis.

Revised version: This study employed four different concentrations of imidacloprid, which were residual concentrations found in honey (0.7 ppb) (Chauzat et al., 2009), bees (1.2 ppb) (Chauzat et al., 2011), pollen (3.1 ppb) (Mullin et al., 2010), and beeswax (377 ppb) (Pareja et al., 2011). The experiment was conducted with three treatment groups: control group (CK), solvent control group (CKac), and imidacloprid-treated group (IMI), each with three replicates. Each replicate contained 32 bees that were distributed in standard 24-well cell culture plates (1 bee/well). The CK group consisted of larval bees fed acetone-royal jelly glucose-fructose-yeast extract-distilled water (0.00377%-50%- 6%-6%-1%-36.99623%, v/w/w/w/w/v). The IMI group consisted of larvae-fed control solution containing imidacloprid (99% active ingredient, Hubei Norna Technology Co. Ltd, Hubei, China) at 0.7 ppb, 1.2 ppb, 3.1 ppb, and 377 ppb. The CKac group served as a solvent control and contained only acetone at a concentration of 0.00377%. A total of 312 larvae were used for survival statistics (96 bees), body weight and width (96 bees), and gene expression, enzyme activity, histochemical content, and histopathologic analysis. Food was administered daily according to age (20 μL for three-day-old larvae, 30 μL for four-day-old larvae, 40 μL for five-day-old larvae, and 50 μL for six-day-old larvae). Dead individuals and food remain were recorded daily, and incubation conditions were 35°C and 96% RH. Surviving larvae were collected 72 h after consumption of the test solution and stored at −80°C for further analysis, including evaluations of gene expression, enzyme activity, and chemical content evaluations. Furthermore, HE-stained sections were used to evaluate any imidacloprid-induced tissue damage and apoptosis.

Developmental phenotypic traits

Larval individuals were randomly selected daily from the CK and IMI groups to measure body weight and width. Food residues, larval mortality, developmental progression, and the amount of time spent in the larval developmental stage (from the first imidacloprid exposure to before pupa) were recorded daily to calculate food consumption, survival rates, developmental rates, and development times. The growth index (GI) was calculated for larvae exposed to imidacloprid from the three-day-old to the six-day-old using the equations described by Zhang (Zhang et al., 1993). The developmental rate was measured as the number of larvae successfully reaching the pupal stage as a percentage of the total initial sample size.

RNA extraction, cDNA synthesis, and real-time quantitative PCR analysis

Earlier version: Total RNA was extracted from larvae using TRIzol reagent (Thermo Fisher Scientific) and RNAex Pro Reagent (Accurate Biology, China) as instructed by the manufacturer. The RNA quality was evaluated by gel electrophoresis and absorbance determination using a NanoDrop 2000 spectrophotometer (Thermo Scientific, New York, USA). Genomic DNA was removed, and cDNA was synthesized using a reverse transcription kit (Evo M-MLV RT Mix Kit with gDNA Clean for qPCR, Accurate Biology, China) according to the manufacturer’s instructions. Briefly, 1 μg of RNA reacted with the 5X gDNA Clean Reaction Mix provided in the kit at 42°C for 5 min to remove residual genomic DNA to ensure the accuracy of the quantitative results. After the reaction, 4 μL of 5X Evo M-MLV RT Reaction Mix and 4 μL of RNase-free water were added to the reaction tube for 20 μL. The mixture was incubated at 37°C for 15 min and 85°C for 5 sec to generate cDNA. Then, quantitative RT–PCR was performed with a Bio-Rad CFX96™ Real-time PCR Detection System (Bio-Rad, Hercules, CA, USA) according to standard protocols and cycling conditions. RP49 (ribosomal protein 49, Accession no. AF441189) (Lourenco et al., 2008) was used as the endogenous control. The primers for target genes and RP49 are listed in Table S1. qPCR system (20 µL): EvaGreen Express 2X qPCR MasterMix (10 µL), each forward and reverse primer (0.4 µL, 0.2 µmol/L), cDNA (1 µL, 100 ng/µL), ddH2O (8.2 µL). The thermal cycling program was as follows: 95°C for 30 sec, followed by 40 cycles of 95°C for 5 s and 60°C for 30 s. The qPCR products were sequenced by Shanghai Biological Engineering Technology Services Co., Ltd. to confirm whether the amplified fragments were the target sequences. Relative mRNA expression levels of target genes were calculated using the 2 method.

Revised version: Total RNA was extracted from larvae using TRIzol reagent (Thermo Fisher Scientific) and RNAex Pro Reagent (Accurate Biology, China) as instructed by the manufacturer. The RNA quality was evaluated by gel electrophoresis and absorbance determination using a NanoDrop 2000 spectrophotometer (Thermo Scientific, New York, USA). Genomic DNA was removed, and cDNA was synthesized using a reverse transcription kit (Evo M-MLV RT Mix Kit with gDNA Clean for qPCR, Accurate Biology, China) according to the manufacturer’s instructions. Briefly, 1 μg of RNA reacted with the 5X gDNA Clean Reaction Mix provided in the kit at 42°C for 5 min to remove residual genomic DNA to ensure the accuracy of the quantitative results. After the reaction, 4 μL of 5X Evo M-MLV RT Reaction Mix and 4 μL of RNase-free water were added to the reaction tube for 20 μL. The mixture was incubated at 37°C for 15 min and 85°C for 5 sec to generate cDNA. Then, quantitative RT–PCR was performed with a Bio-Rad CFX96™ Real-time PCR Detection System (Bio-Rad, Hercules, CA, USA) according to standard protocols and cycling conditions. RP49 (ribosomal protein 49, Accession no. AF441189) (Lourenco et al., 2008) was used as the endogenous control. The primers for target genes and RP49 are listed in Table S1. qPCR system (20 µL): EvaGreen Express 2X qPCR MasterMix (10 µL), each forward and reverse primer (0.4 µL, 0.2 µmol/L), cDNA (1 µL, 100 ng/µL), ddH2O (8.2 µL). The thermal cycling program was as follows: 95°C for 30 sec, followed by 40 cycles of 95°C for 5 s and 60°C for 30 s. The qPCR products were sequenced by Shanghai Biological Engineering Technology Services Co., Ltd. to confirm whether the amplified fragments were the target sequences. Relative mRNA expression levels of target genes were calculated using the 2−ΔΔCt method.

Histochemical measurement

Juvenile hormone 3 (JH-III), 20-hydroxyecdysone (20E) titers, and reactive oxygen species (ROS) levels in larvae were determined using ELISA enzyme-linked immunosorbent assay kits (Jining, Shanghai, China) according to the manufacturer’s instructions. Briefly, the larva was homogenized on ice in 5 volumes (M: V) of phosphate-buffered saline (PBS) buffer, and the supernatant was collected after centrifugation at 12,000 x g for 20 min. Subsequently, 50 μL of reference standard solution or 10 μL of sample supernatant plus 40 μL of dilution were added to the microtiter wells and incubated at 37°C for 60 min with 100 μL of HRP-labeled detection antibody. Each well was washed five times with 400 μL of PBS buffer containing 0.05% Tween 20, pat dried, and left to stand at room temperature for 5 min. TMB chromogen solution A (50 μL) and solution B (50 μL) were added to each well, gently mixed, and incubated at 37°C for 15 min in the dark. Then, 50 μL of H2SO4 (4 mol/L) was added to each well to terminate the reaction. Finally, the OD value was measured at 450 nm with a NanoDrop 2000 spectrophotometer (Thermo Scientific, New York, USA), and the content of the detected substance was calculated from a calibration curve constructed by reference standards.

Earlier version: To prepare samples for analysis of protein carbonylation (PCO), acetylcholinesterase (AchE), catalase (CAT), malondialdehyde (MDA), total antioxidant capacity (T-AOC), superoxide dismutase (SOD), and glutathione (GSH), larvae was weighed and homogenized on ice in 5 volumes (M: V) of PBS buffer, then centrifuged at 12,000 x g for 15 min to collect the supernatant. Subsequently, PCO content was measured using the 2,4-dinitrophenyl-hydrazine (DNPH) method with slight modifications. The above-prepared supernatant was mixed with streptomycin sulfate at a 9:1 volume (V: M), kept at room temperature for 10 min, and then centrifuged at 12,000 x g for 15 min. Then, 80 μL of supernatant was mixed with 160 μL of 10LrmM DNPH in 2LrM HCl and reacted at 37°C for 30 min. Sample controls were also prepared by adding an equal volume (500LrμL) of 2LrM HCl. Then, 200 μL of 20% TCA was added, mixed thoroughly, and centrifuged at 4°C for 10 min at 12,000 x g, and the precipitate was collected. Add 400 μL of ethanol: ethyl acetate mixture (1:1), vortex and mix, centrifuge at 4°C 12,000 x g for 10 min, and collect the precipitate to remove unbound DNPH. This DNPH removal step was repeated once. The pellet was mixed with 400LrμL of 6LrM guanidine hydrochloride, kept at room temperature for 15 min, and then centrifuged at 4°C 12,000 x g for 10 min to collect the supernatant. The carbonyl and protein contents were determined by reading the absorbance at λLr=Lr370Lrnm and 280 nm with a SpectraMax 190 spectrophotometer (Molecular Devices, California, USA) for each sample against an appropriate blank. Finally, the carbonyl content per unit of protein was calculated based on the protein content.

Revised version: To prepare samples for analysis of protein carbonylation (PCO), acetylcholinesterase (AchE), catalase (CAT), malondialdehyde (MDA), total antioxidant capacity (T-AOC), superoxide dismutase (SOD), and glutathione (GSH), larvae was weighed and homogenized on ice in 5 volumes (M: V) of PBS buffer, then centrifuged at 12,000 x g for 15 min to collect the supernatant. Subsequently, PCO content was measured using the 2,4-dinitrophenyl-hydrazine (DNPH) method with slight modifications. The above-prepared supernatant was mixed with streptomycin sulfate at a 9:1 volume (V: M), kept at room temperature for 10 min, and then centrifuged at 12,000 x g for 15 min. Then, 80 μL of supernatant was mixed with 160 μL of 10 mM DNPH in 2 M HCl and reacted at 37°C for 30 min. Sample controls were also prepared by adding an equal volume (500 μL) of 2 M HCl. Then, 200 μL of 20% TCA was added, mixed thoroughly, and centrifuged at 4°C for 10 min at 12,000 x g, and the precipitate was collected. Add 400 μL of ethanol: ethyl acetate mixture (1:1), vortex and mix, centrifuge at 4°C 12,000 x g for 10 min, and collect the precipitate to remove unbound DNPH. This DNPH removal step was repeated once. The pellet was mixed with 400 μL of 6 M guanidine hydrochloride, kept at room temperature for 15 min, and then centrifuged at 4°C 12,000 x g for 10 min to collect the supernatant. The carbonyl and protein contents were determined by reading the absorbance at λ = 370 nm and 280 nm with a SpectraMax 190 spectrophotometer (Molecular Devices, California, USA) for each sample against an appropriate blank. Finally, the carbonyl content per unit of protein was calculated based on the protein content.

Earlier version: AChE activity was measured by the Ellman method as described by Orhan et al. (Orhan et al., 2007) with slight modifications. Briefly, 20 μL of the prepared sample supernatant was mixed with 160 µL of 0.1 mM PBS buffer (pH 8.0) and 10 µL of an AChE (0.28 U mL) and incubated for 15 min at room temperature. Then, 10 µL of ATChI (acetylthiocholine iodide, 4 mg mL) was added together with 20 µL of DTNB (5,5-dithiobis-2-nitrobenzoic acid, 1.2 mg mL) and immediately mixed, and the absorbance was read at 412 nm. The results were expressed as activity percentages relative to the sample solvent negative control.

Revised version: AChE activity was measured by the Ellman method as described by Orhan et al. (Orhan et al., 2007) with slight modifications. Briefly, 20 μL of the prepared sample supernatant was mixed with 160 µL of 0.1 mM PBS buffer (pH 8.0) and 10 µL of an AChE (0.28 U mL−1) and incubated for 15 min at room temperature. Then, 10 µL of ATChI (acetylthiocholine iodide, 4 mg mL−1) was added together with 20 µL of DTNB (5,5- dithiobis-2-nitrobenzoic acid, 1.2 mg mL−1) and immediately mixed, and the absorbance was read at 412 nm. The results were expressed as activity percentages relative to the sample solvent negative control.

Earlier version: The CAT activity was determined using a kit (Beyotime, Shanghai, China) according to the manufacturer’s instructions with slight modifications. Briefly, 10 μL of the prepared sample supernatant was mixed with 30 μL of Tris-HCl buffer (pH 7.0, 50 mmol L), 10 µL of 250 mM H2O2 was added, and the mixture was immediately mixed and kept at 25°C for 5 min. Then, 450 μL of 1.8 mol L-1 H2SO4 was added, and the mixture was vortexed to terminate the reaction. Subsequently, 10 µL of this termination reaction solution was mixed with 40 µL of Tris-HCl buffer (pH 7.0, 50 mmol L). Then, 10 µL of this mixture was added to a 96-well plate mixed with 200 µL of color development working solution containing peroxidase and incubated for 15 min at 25°C. Finally, the absorbance was measured at 520 nm, and the CAT activity was calculated using a calibration curve of 5 mM H2O2.

Revised version: The CAT activity was determined using a kit (Beyotime, Shanghai, China) according to the manufacturer’s instructions with slight modifications. Briefly, 10 μL of the prepared sample supernatant was mixed with 30 μL of Tris-HCl buffer (pH 7.0, 50 mmol L-1), 10 µL of 250 mM H2O2 was added, and the mixture was immediately mixed and kept at 25°C for 5 min. Then, 450 μL of 1.8 mol L-1 H2SO4 was added, and the mixture was vortexed to terminate the reaction. Subsequently, 10 µL of this termination reaction solution was mixed with 40 µL of Tris-HCl buffer (pH 7.0, 50 mmol L-1). Then, 10 µL of this mixture was added to a 96-well plate mixed with 200 µL of color development working solution containing peroxidase and incubated for 15 min at 25°C. Finally, the absorbance was measured at 520 nm, and the CAT activity was calculated using a calibration curve of 5 mM H2O2.

T-AOC was measured using the total antioxidant capacity assay kit with the ABTS (2,2’-azino-bis(3-ethylbenzthiazoline-6-sulfonic acid) method based on a previous report (Chen et al., 2019). Add 10 μL of the prepared sample supernatant to 200 μL of ABTS working solution. After 6 min of reaction, the absorbance was measured at 734 nm. Finally, the antioxidant activity was displayed by Trolox equivalents antioxidant capacity as mmol Trolox equivalents/g extract. The ABTS working solution was prepared as follows: 400 μL of potassium persulfate was added to 400 μL of ABTS to generate an ABTS stock solution, which was then incubated for 16 h at room temperature in the dark. Before use, the stock solution was diluted with phosphate buffer, and the absorbance at 734 nm was adjusted to 0.70 ± 0.05.

Additionally, MDA, SOD, and GSH levels were detected using ELISA enzyme-linked immunosorbent assay kits (Beyotime, Shanghai, China) according to our recent study (Li et al., 2022).

Histopathological analysis

Microsectioning and hematoxylin-eosin (HE) staining was performed according to standard histological protocols (Khan et al., 2010) with some modifications. Briefly, entire fresh larvae were fixed in 10% neutral-buffered formalin solution for 12 h at 4°C after 72 h of imidacloprid exposure. After isopropyl alcohol dehydration, samples were embedded in paraffin and sectioned. Three-micron-thick sections were stained with a Hematoxylin and Eosin Staining Kit (Beyotime, Shanghai, China), observed with an Olympus BX51 microscope, and photographed with an Olympus digital camera.

Statistical analysis

Data were obtained from three independent experiments performed in triplicate and are presented as the mean ± SEM. SPSS statistical package was used for statistical analysis using Student’s t-test and one-way ANOVA followed by Tukey’s multiple comparison test (*P < 0.05, **P < 0.01, and ***P < 0.001 were considered statistically significant). GraphPad software was used for data normalization, Pearson correlation coefficient calculation was analyzed using SPSS, and a heat map was generated.

Supporting information